---
title: "Visualizations with ggplot2"
author: "Andrew Ba Tran"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---


<style>
.cell-left {
  text-align: left;
}

.cell-right {
  text-align: right;
}

.cell-center {
  text-align: center;
}

table {
    margin: auto;
    border-top: 1px solid rgb(102, 102, 102);
    border-bottom: 1px solid rgb(102, 102, 102);
    display: table;
    border-collapse: separate;
    box-sizing: border-box;
    border-spacing: 2px;
    border-color: grey;
    padding-bottom:5px;
}

</style>


```{r setup, include=FALSE, echo=FALSE}
library(tidyverse)
library(learnr)
library(lubridate)
library(readxl)
library(DT)
library(flair)
library(janitor)
library(gradethis)


sd <- read_csv("data/san_diego.csv")

sd <- clean_names(sd)


sd_adjusted <- sd %>% 
  mutate(death_date=mdy(death_date)) %>% 
  mutate(month=month(death_date, label=TRUE, abbr=TRUE)) 

sd_month <- sd_adjusted %>% 
  count(year, month, manner_of_death, name="deaths") %>% 
  mutate(date=mdy(paste0(month, " 1, ", year)))

scented <- read_excel("data/Scented_all.xlsx", sheet=1)
unscented <- read_excel("data/Unscented_all.xlsx", sheet=1)

scented <- scented %>% 
  mutate(type="scented")
  
unscented <- unscented %>% 
  mutate(type="unscented")

scented <- bind_rows(scented, unscented)

scented_summary <- scented %>%
  arrange(Date) %>%
  filter(Date >= "2017-01-01") %>%
  mutate(Date=ymd(Date)) %>% 
  filter(CandleID <= 3) %>%
  group_by(Date, type) %>%
  summarise(Average_Rating=mean(Rating),
            Reviews=n())
  
sd_year <- sd_adjusted %>% 
  # if we're going to compare this year to previous years, we need to exclude data we don't have yet
  filter(!month %in% c("Oct", "Nov", "Dec")) %>% 
  count(year, manner_of_death, name="deaths") 


no_scent_analysis <- scented %>% 
  # filter the reviews so we're only looking at scented candles
  filter(type=="scented") %>% 
  arrange(Date) %>% 
  # only focus on the reviews from this year
  filter(Date >= mdy("01-01-2020")) %>% 
  # create a new column no_scent that will fill a 1 if No Scent is mentioned
  mutate(no_scent=case_when(
    str_detect(Review, "[Nn]o scent") ~ 1, 
    str_detect(Review, "[Nn]o smell") ~ 1,
    str_detect(Review, "[Dd]oes not smell like") ~ 1,
    str_detect(Review, "[Dd]oesn't smell like") ~ 1,
    str_detect(Review, "[Cc]an't smell") ~ 1,
    str_detect(Review, "[Cc]annot smell") ~ 1,
    str_detect(Review, "[Ff]aint smell") ~ 1,
    str_detect(Review, "[Ff]aint scent") ~ 1,
    str_detect(Review, "[Dd]on't smell") ~ 1,
    str_detect(Review, "[Ll]ike nothing") ~ 1,
    TRUE ~ 0
  )) %>% 
  # create a new month column
  mutate(month=month(Date, label=TRUE)) %>% 
  # let's aggregate around month
  group_by(month) %>% 
  # we just need to add up the 1s in no_scent and count up the total reviews with n()
  summarize(no_scent=sum(no_scent), reviews=n()) %>% 
  # and do some math
  mutate(percent=no_scent/reviews*100)
```

Note: All answers to exercises can be found [here](https://github.com/utdata/chj-r-introduction/blob/main/chj-day4/chj-day4-answers.Rmd).

## Data introduction

Before we begin, let's bring in some data. 

Download these files to your project folder.

```{r download}
# https://github.com/kateptrv/Candles/raw/main/Scented_all.xlsx
# https://github.com/kateptrv/Candles/raw/main/Unscented_all.xlsx
```

These are Excel files and not CSVs, so we need a special package to import those tricky files in. We'll use one called **readxl**.

```{r import_fake, eval=F}
library(readxl)

scented <- read_excel("data/Scented_all.xlsx", sheet=1)
unscented <- read_excel("data/Unscented_all.xlsx", sheet=1)
```

Let's take a look at what were working with:

```{r view_unscented}
library(DT)
library(dplyr)

scented <- read_excel("data/Scented_all.xlsx", sheet=1)
unscented <- read_excel("data/Unscented_all.xlsx", sheet=1)

scented <- scented %>% 
  mutate(type="scented")
  
unscented <- unscented %>% 
  mutate(type="unscented")

scented <- bind_rows(scented, unscented)

scented %>% 
  head(4) %>% 
  datatable()
```

These are more than 21,000 reviews for the top 5 rated scented and unscented candles on Amazon. It goes back years.

Let's quickly filter this data to the top 3 candles only and figure out the `Average_Rating` for each day. And also count the total `Reviews` for each day. And we should convert the dates in the `Date` column with a **lubridate** function.

```{r lub, echo=F}
library(lubridate)
```

### Exercise 1

Fill in the blanks with the right code.

```{r scented, exercise=TRUE}
scented_summary <- scented %>%
  arrange(Date) %>%
  filter(Date >= "2017-01-01") %>%
  # convert the String Date into a date format using lubridate function
  mutate(Date=___(Date)) %>% 
  filter(CandleID <= 3) %>%
  group_by(Date, type) %>% 
  # Now how do we aggregate average and count?
  _________(Average_Rating=____(Rating),
            Reviews=___)
  
scented_summary
```


```{r scented-solution}
scented_summary <- scented %>%
  arrange(Date) %>%
  filter(Date >= "2017-01-01") %>%
  # convert the String Date into a date format using lubridate function
  mutate(Date=ymd(Date)) %>% 
  filter(CandleID <= 3) %>%
  group_by(Date, type) %>%
  # Now how do we aggregate average and count?
  summarise(Average_Rating=mean(Rating),
            Reviews=n())

scented_summary
```

```{r scented-code-check}
grade_code()
```

<div id="scented-hint">
**Hint:** ymd(), group_by(), summarize(), mean(), n()
</div>

We've narrowed down 21,000 rows of data to a more workable 2,471.

This summarized data set will be our entry to the:

## Grammar of Graphics

Here's a quick chart. You may not understand the code yet, but I'll break it down for you after we talk about some concepts.

```{r first_chart}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, colour=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month")
```

### Mapping data to aesthetics

**Aesthetic**

* Visual property of a graph

* Position, shape, color, etc.

**Data**

* A column in a dataset

-----

Here's are the core components of this chart:

<table>
  <tr>
    <th class="cell-left">Data</th>
    <th class="cell-left">Aesthetic</th>
    <th class="cell-left">Graphic/Geometry</th>
  </tr>
  <tr>
    <td class="cell-left">Dates</td>
    <td class="cell-left">Position (x-axis)&emsp;</td>
    <td class="cell-left">Point</td>
  </tr>
  <tr>
    <td class="cell-left">Rating</td>
    <td class="cell-left">Position (y-axis)</td>
    <td class="cell-left">Point</td>
  </tr>
    <tr>
    <td class="cell-left">Candle type</td>
    <td class="cell-left">Color</td>
    <td class="cell-left">Point</td>
  </tr>
  <tr>
    <td class="cell-left">Single date</td>
    <td class="cell-left">Position (x-intercept)</td>
    <td class="cell-left">Line</td>
  </tr>
</table>

-----

Here's how it's interpreted in ggplot2 code

<table>
  <tr>
    <th class="cell-left">Data</th>
    <th class="cell-left">aes()</th>
    <th class="cell-left">geom</th>
  </tr>
  <tr>
    <td class="cell-left">Dates</td>
    <td class="cell-left">x</td>
    <td class="cell-left">geom_point()</td>
  </tr>
  <tr>
    <td class="cell-left">Average_Rating</td>
    <td class="cell-left">y</td>
    <td class="cell-left">geom_point()</td>
  </tr>
  <tr>
    <td class="cell-left">type</td>
    <td class="cell-left">color</td>
    <td class="cell-left">geom_point()</td>
  </tr>
  <tr>
    <td class="cell-left">1/20/2020</td>
    <td class="cell-left">xintercept</td>
    <td class="cell-left">geom_vline()</td>
  </tr>
</table>



### ggplot() template

Here's the dataframe called `scented_summary` as a reminder:

```{r preview}
scented_summary %>% head(5)
```


```{r show-ggplot-template, echo=FALSE, tidy=FALSE}
decorate('
ggplot(data = DATA) +
  GEOM_FUNCTION(mapping = aes(AESTHETIC MAPPINGS))
', eval = FALSE) %>% 
  flair("DATA", background = "#CBB5FF", before = "<b>", after = "</b>") %>% 
  flair("GEOM_FUNCTION", background = "#FFDFD1", before = "<b>", after = "</b>") %>% 
  flair("AESTHETIC MAPPINGS", background = "#FFD0CF", before = "<b>", after = "</b>") %>% 
  knit_print.with_flair()
```

--

```{r ggplot-template-example, echo=FALSE}
decorate('
ggplot(data = scented_summary) +
  geom_point(mapping = aes(x = Date, y = Average_Rating, color = type)) +
  geom_vline(xintercept=ymd("2020-01-20")) 
', eval = FALSE) %>%
  flair("scented_summary", background = "#CBB5FF", before = "<b>", after = "</b>") %>% 
  flair("geom_point", background = "#FFDFD1", before = "<b>", after = "</b>") %>% 
  flair("geom_vline", background = "#FFDFD1", before = "<b>", after = "</b>") %>% 
  flair_rx('x = Date|y = Average_Rating|color = type|xintercept= ymd("2020-01-20")', 
           background = "#FFD0CF", before = "<b>", after = "</b>") %>% 
  knit_print.with_flair()
```

---

```{r img0, echo = F, out.width="100%"}
library(knitr)
include_graphics("images/ggplot1.png")
``` 

### Grammatical layers

So far we know about data, aesthetics, and geometries.

Think of these components as **layers**.



```{r img1a, echo = F, out.width="50%"}
include_graphics("images/ggplot1a.png")
``` 

Add them to foundational `ggplot()` with +

**Possible aesthetics**

```{r img2a, echo = F, out.width="100%"}
include_graphics("images/ggplot2a.png")
``` 


**Possible geoms**

```{r img3a, echo = F, out.width="60%"}
include_graphics("images/ggplot3a.png")
``` 

THERE ARE [SO MANY GEOMS](https://twitter.com/search?q=geom_%20%23rstats&src=typed_query&f=image) for different visualizations. Here are the [official ones](https://ggplot2.tidyverse.org/reference/index.html#section-layer-geoms).

### Additional layers

There are many of other grammatical layers we can use to describe graphs.

We sequentially add layers onto the foundational `ggplot()` plot to create complex figures.


```{r img4a, echo = F, out.width="50%"}
include_graphics("images/ggplot4a.png")
``` 

Scales change the properties of the variable mapping.

<table>
  <tr>
    <th class="cell-left">Example layer</th>
    <th class="cell-left">What it does</th>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">scale_x_continuous()</code></td>
    <td class="cell-left">Make the x-axis continuous</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">scale_x_continuous(breaks = 1:5)&ensp;</code></td>
    <td class="cell-left">Manually specify axis ticks</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">scale_x_date()</code></td>
    <td class="cell-left">Considers x-axis dates</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">scale_color_gradient()</code></td>
    <td class="cell-left">Use a gradient</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">scale_fill_viridis_d()</code></td>
    <td class="cell-left">Fill with discrete viridis colors</td>
  </tr>
</table>

----

Look at the difference between these two charts without and with a scales function. 

```{r second_chart}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20"))
```

Check out the x-axis.

### Exercise 2

Now add `scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month")` to the bottom of the code


```{r scented2, exercise=TRUE}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +

# Add line above this one
```

```{r scented2-solution}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month")
```

```{r scented2-code-check}
grade_code()
```

Do you see the difference at the bottom of the chart compared to the one above it?

### Facets

Facets show subplots for different subsets of data.

<table>
  <tr>
    <th class="cell-left">Example layer</th>
    <th class="cell-left">What it does</th>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">facet_wrap(vars(type))</code></td>
    <td class="cell-left">Plot for each candle type</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">facet_wrap(vars(type, year))</code>&emsp;</td>
    <td class="cell-left">Plot for each candle type/year</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">facet_wrap(..., ncol = 1)</code></td>
    <td class="cell-left">Put all facets in one column</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">facet_wrap(..., nrow = 1)</code></td>
    <td class="cell-left">Put all facets in one row</td>
  </tr>
</table>


```{r facet_example1, warning=F, message=F}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month") +
  facet_wrap(vars(type))
```

### Exercise 3

Now, try it with nrow=1 or ncol=1

```{r facet_example2, exercise=TRUE}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month") +
  facet_wrap(vars(type), ___ = 1)
```


```{r facet_example2-solution}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month") +
  facet_wrap(vars(type), ___ = 1)
```


```{r facet_example2-code-check}
grade_code()
```


### Labels


<table>
  <tr>
    <th class="cell-left">Example layer</th>
    <th class="cell-left">What it does</th>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">labs(title = "Neat title")</code></td>
    <td class="cell-left">Title</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">labs(caption = "Something")</td>
    <td class="cell-left">Caption</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">labs(y = "Something")</td>
    <td class="cell-left">y-axis</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">labs(color = "Type")</code></td>
    <td class="cell-left">Title of size legend</td>
  </tr>
</table>

```{r labels-example, tidy=FALSE, message=FALSE}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month") +
  labs(title = "Candle reviews on Amazon.com",
       subtitle = "Top 3 most popular candles 2017-2020",
       x = "Date",
       y = "Average daily rating (1-5)",
       color = "Type",
       caption = "Source: Kate Petrova")
```

### Theme

Change the appearance of anything in the plot.

There are many built-in themes.


<table>
  <tr>
    <th class="cell-left">Example layer</th>
    <th class="cell-left">What it does</th>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">theme_grey()</code></td>
    <td class="cell-left">Default grey background</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">theme_bw()</td>
    <td class="cell-left">Black and white</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">theme_dark()</td>
    <td class="cell-left">Dark</td>
  </tr>
  <tr>
    <td class="cell-left"><code class="remark-inline-code">theme_minimal()</code></td>
    <td class="cell-left">Minimal</td>
  </tr>
</table>

### Exercise 4

Try out the different themes listed above in the code below.

```{r themes-example, exercise=TRUE, tidy=FALSE, message=FALSE}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month") +
  labs(title = "Candle reviews on Amazon.com",
       subtitle = "Top 3 most popular candles",
       x = "Date",
       y = "Average rating",
       color = "Type",
       caption = "Source: Kate Petrova") +
  theme______
```

### More themes

There are a collections of pre-built themes online, like the [ggthemes package](https://jrnold.github.io/ggthemes/).

Organizations often make their own custom themes, like [the BBC](https://bbc.github.io/rcookbook/).


```{r img10, echo = F, out.width="100%"}
include_graphics("images/bbc-cookbook.png")
``` 

### Theme adjustments

Make theme adjustments with `theme()`

There are a billion options here!

Add this chunk of code in the exercise below it:

```{r eval=F}
theme_bw() + 
theme(legend.position = "bottom",
      plot.title = element_text(face = "bold"),
      panel.grid = element_blank(),
      axis.title.y = element_text(face = "italic"))
```

### Exercise 5

```{r themes_example2, exercise=TRUE}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month") +
  labs(title = "Candle reviews on Amazon.com",
       subtitle = "Top 3 most popular candles",
       x = "Date",
       y = "Average rating",
       color = "Type",
       caption = "Source: Kate Petrova") +
  theme______

# copy and paste the chunk of code above this chunk into the line above this one
```


```{r themes_example2-solution}
ggplot(data= scented_summary) +
  geom_point(mapping=aes(x= Date, y= Average_Rating, color=type)) +
  geom_vline(xintercept= ymd("2020-01-20")) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "6 month") +
  labs(title = "Candle reviews on Amazon.com",
       subtitle = "Top 3 most popular candles",
       x = "Date",
       y = "Average rating",
       color = "Type",
       caption = "Source: Kate Petrova") +
  theme_bw() + 
  theme(legend.position = "bottom",
        plot.title = element_text(face = "bold"),
        panel.grid = element_blank(),
        axis.title.y = element_text(face = "italic"))
```

```{r themes_example2-code-check}
grade_code()
```


These were just a few examples of layers.

```{r img11, echo = F, out.width="50%"}
include_graphics("images/ggplot5.png")
``` 

See the [ggplot2 documentation](https://ggplot2.tidyverse.org/reference/index.html) for complete examples of everything you can do

One last thing: I want to add one more geom_ so that it's more obvious what the pattern is.

`geom_smooth()`

I'm also going to move the mapping to the first line so it will be consistent for all `geoms_` after.

```{r themes-example-last, tidy=FALSE, message=FALSE, fig.width=8}
ggplot(data= scented_summary, 
       mapping=aes(x= Date, y= Average_Rating)) +
  geom_point(alpha = 0.2, colour = "goldenrod3") +
geom_smooth(method = "loess", size = 1.5, colour = "goldenrod3", fill = "goldenrod3") +
  geom_vline(xintercept= ymd("2020-01-20")) +
  facet_wrap(vars(type), ncol=2) +
  scale_x_date(date_labels = "%m-%Y", date_breaks = "12 month") +
  labs(title = "Candle reviews on Amazon.com",
       subtitle = "Top 3 most popular candles",
       x = "Date",
       y = "Average rating",
       color = "Type",
       caption = "Source: Kate Petrova") +
  theme_minimal()
```

 It seems obvious that there's a trend here, right?
 
 Why do you think that is?
 
 ###Let's find out
 
Let's add a column flagging whether any variation of "no scent" is mentioned in the review and figure out the percent of those complaints in each month of complaints. We're going to use `case_when()` and `str_detect()`!
 
 Here's the code for you with comments to explain what's going on.
 
```{r no_scent, warning=F, message=F}
 
no_scent_analysis <- scented %>% 
  # filter the reviews so we're only looking at scented candles
  filter(type=="scented") %>% 
  arrange(Date) %>% 
  # only focus on the reviews from this year
  filter(Date >= mdy("01-01-2020")) %>% 
  # create a new column no_scent that will fill a 1 if No Scent is mentioned
  mutate(no_scent=case_when(
    str_detect(Review, "[Nn]o scent") ~ 1, 
    str_detect(Review, "[Nn]o smell") ~ 1,
    str_detect(Review, "[Dd]oes not smell like") ~ 1,
    str_detect(Review, "[Dd]oesn't smell like") ~ 1,
    str_detect(Review, "[Cc]an't smell") ~ 1,
    str_detect(Review, "[Cc]annot smell") ~ 1,
    str_detect(Review, "[Ff]aint smell") ~ 1,
    str_detect(Review, "[Ff]aint scent") ~ 1,
    str_detect(Review, "[Dd]on't smell") ~ 1,
    str_detect(Review, "[Ll]ike nothing") ~ 1,
    TRUE ~ 0
  )) %>% 
  # create a new month column
  mutate(month=month(Date, label=TRUE)) %>% 
  # let's aggregate around month
  group_by(month) %>% 
  # we just need to add up the 1s in no_scent and count up the total reviews with n()
  summarize(no_scent=sum(no_scent), reviews=n()) %>% 
  # and do some math
  mutate(percent=no_scent/reviews*100)

no_scent_analysis
```

**Make me a chart of your choice using the data above.**

(Actually, maybe not geom_line()-- that one doesn't work on this data for some reason)

### Exercise 6
   
```{r no-smell-example2, exercise=TRUE}
## whatever you like (okay, I prefer if you use the percent column)
ggplot(no_scent_analysis) +

```
